// variables1.rs
// Make me compile! Execute the command `rustlings hint variables1` if you want a hint :)

// About this `I AM NOT DONE` thing:
// We sometimes encourage you to keep trying things on a given exercise,
// even after you already figured it out. If you got everything working and
// feel ready for the next exercise, remove the `I AM NOT DONE` comment below.

fn main() {
    // let x = 5;
    // let x : u32 = 5;
    // let x : u64 = 5;
    // let x : u128 = 5;

    // Error: no type i4 let x : i4 = 5;
    // let x : i8 = 5;
    // let x : i16 = 5;
    // let x : i32 = 5;
    // let x : i64 = 5;
    let x : i128 = 5;

    println!("x has the value {}", x);
}
